# COGS 18 - Exam 2 (Take-Home)

This is the take-home portion of the second exam for Spring 2024. It covers through the Classes lecture.

This take-home portion of the exam is out of 2.5 points, worth 2.5% of your grade.

**PLEASE DO NOT CHANGE THE NAME OF THIS FILE.**

**PLEASE DO NOT COPY & PASTE OR DELETE CELLS INLCUDED IN THE ASSIGNMENT.** (Note that you can add additional cells, if you want to test things out.

**In-class Answer Key**

- 11AM [here](https://docs.google.com/document/d/1CUHJ-HJy_n78vjNXdfwEdxlf_kOsc0nyDDaTnxHZwec/edit?usp=sharing)
- 2PM [here](https://docs.google.com/document/d/1vV-Obolcu58zdycexAldY8xHokgeWekMyE0HXnEpiFM/edit?usp=sharing)

#### Instructions

#### Timing
- There take-home portion will be released 3:30 the day of the exam and due the same day at 11:59PM

#### The Rules
- You are to complete this exam on your own.
- This is open-notes & open-Internet.
- You may not talk to any humans about this exam. 
- Clarification questions will ***not*** be allowed and there will be no posting on Piazza about the exam at all. 
     - If you are confused about wording, add a note to your exam explaining your confusion and how you interpreted the question. 
    - If you have a technical issue completing the exam, message (Piazza or email) Prof Ellis ASAP.

 <span style="color: red;">Note: </span> This portion of the midterm will be autograded, similar to an assignment. `assert` statements are there to guide your thinking but do not guarantee you've gotten the answer correct.

### Q0 - Honor Code (0.00 points)

In the cell below, include a variable `honor_code` that stores the boolean `True` if you agree to the following statement:

>I agree that this exam was completed individually with the knowlege in my brain, information the notes from this course, and/or with searching for help on the Internet *without searching for answers to the text from these questions directly*. I did not ask anyone about specific questions on this exam. I did not post these questions (in part or in whole) on the Internet. I did not copy answers to these questions from anywhere or anyone else. I understand all code I've written on this exam and could explain my answers to someone else if asked.


<!-- grade_id: cell-9dae831891d7a1ad, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
honor_code = True
### END SOLUTION
```

<!-- grade_id: cell-8a6d1a928512e45c, points: 0, grade: true, solution: false -->
```python
assert honor_code
```

### Background

On the in-class exam you saw the following:

> The idea here is that there will be a single class (`Rectangle`) with a few attributes (described below) and two methods: 1) `calculate_area()` and  2) `calculate_perimeter()`. 

In the next two questions, you'll write the code and create some `Rectangle`s!

### Q1 - `Rectangle` (1.5 pts)

Here, create a `Rectangle()` type object, with the following:

- instance attributes: 1) `length`  and 2) `width`, each of which the user will have to provide when an instance of the object is first created.
- methods:
    - `calculate_area()`, which should calculate and `return` the rectangle’s area. Note that a rectangle’s area is its length times its width
    - `calculate_perimeter()` , which should calculate and return the rectangle’s perimeter. Note that perimeter is calculated by adding the length of every side of the rectangle together

Checking your work: `rect = Rectangle(length=4, width=5)` should initialize a `Rectangle` type object. After initialization, calling `calculate_area()` on `rect` should `return` 20. Calling `calculate_perimeter()` should `return` 18. 

<!-- grade_id: cell-890ee14a107fa8f5, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
class Rectangle():
    def __init__(self, length, width):
        self.length = length
        self.width = width

    def calculate_area(self):
        return self.length * self.width

    def calculate_perimeter(self):
        return 2 * (self.length + self.width)
### END SOLUTION
```

```python
# you can use this cell to test/execute/check your thinking (optional)
```

<!-- grade_id: cell-986ea16845831923, points: 0.25, grade: true, solution: false -->
```python
# check no typo in class name
assert callable(Rectangle)
```

<!-- grade_id: cell-d9a47a609cfdb98b, points: 0.25, grade: true, solution: false -->
```python
# hidden test to check that Rectangle can be initialized
### BEGIN HIDDEN TESTS
my_rectangle = Rectangle(length=4, width=5)
my_rectangle2 = Rectangle(length=2, width=3)
### END HIDDEN TESTS
```

<!-- grade_id: cell-0ad5a91264167d53, points: 0.25, grade: true, solution: false -->
```python
# hidden test that check length instance attribute
### BEGIN HIDDEN TESTS
assert my_rectangle.length == 4
assert my_rectangle2.length == 2
### END HIDDEN TESTS
```

<!-- grade_id: cell-bf0ee3500fcd2bc9, points: 0.25, grade: true, solution: false -->
```python
# hidden test that check width instance attribute
### BEGIN HIDDEN TESTS
assert my_rectangle.width == 5
assert my_rectangle2.width == 3
### END HIDDEN TESTS
```

<!-- grade_id: cell-7d903b9bfc07d385, points: 0.25, grade: true, solution: false -->
```python
# hidden tests that check calculate_area method
### BEGIN HIDDEN TESTS
assert my_rectangle.calculate_area() == 20
assert my_rectangle2.calculate_area() == 6
### END HIDDEN TESTS
```

Output:
```
True
```

<!-- grade_id: cell-0693f283f6eaafeb, points: 0.25, grade: true, solution: false -->
```python
# hidden tests that check calculate_perimeter method
### BEGIN HIDDEN TESTS
assert my_rectangle.calculate_perimeter() == 18
assert my_rectangle2.calculate_perimeter() == 10
### END HIDDEN TESTS
```

Output:
```
True
```

<!-- grade_id: cell-5d4daa1f4b374aa1, points: 0, grade: true, solution: false -->
```python
# Do nothing with this cell
# This cell here ensures that even if you got Q1 wrong 
# but got Q2 right, you'll get credit for Q2
### BEGIN HIDDEN TESTS
# this ensures you have the correct Rectangle class is used 
class Rectangle():
    def __init__(self, length, width):
        self.length = length
        self.width = width

    def calculate_area(self):
        return self.length * self.width

    def calculate_perimeter(self):
        return 2 * (self.length + self.width)
### END HIDDEN TESTS
```

### Q2 - Using your class (1 pt)

Here, you’ll initialize `rect1` and `rect2`, two different instances of your `Rectangle` object.

- `rect1` should have a length of 10 and a width of 5. 
- `rect2` should have a length of 3 and a width of 9. 

After initializing each `Rectangle` object, let's utilize your methods:

- Call the `calculate_area()` method on `rect1`, storing the output in `rect1_area`
- Call the `calculate_perimeter()` method on `rect2`, storing the output in `rect2_perimeter`


<!-- grade_id: cell-5406a3284790c7de, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
rect1 = Rectangle(length=10, width=5)
rect2 = Rectangle(length=3, width=9)

rect1_area = rect1.calculate_area()
rect2_perimeter = rect2.calculate_perimeter()
### END SOLUTION
```

```python
# you can use this cell to test/execute/check your thinking (optional)
```

<!-- grade_id: cell-b6f17c05f4455dff, points: 0.25, grade: true, solution: false -->
```python
# check for no typos in Rectangle objects
assert rect1
assert rect2
```

<!-- grade_id: cell-3ed4c4555bf2aad6, points: 0.25, grade: true, solution: false -->
```python
# hidden test that checks rect1 and rect2 are of the expected type
### BEGIN HIDDEN TESTS
assert type(rect1) == Rectangle
assert type(rect2) == Rectangle
### END HIDDEN TESTS
```

<!-- grade_id: cell-1ce6db55cfb0fe1a, points: 0.25, grade: true, solution: false -->
```python
# check for typos in variable name
assert rect1_area

# hidden test to check for correct value in rect1_area
### BEGIN HIDDEN TESTS
assert rect1_area == 50
### END HIDDEN TESTS
```

<!-- grade_id: cell-1ff39fd3a0967d23, points: 0.25, grade: true, solution: false -->
```python
# check for typos in variable name
assert rect2_perimeter

# hidden test to check for correct value in rect2_perimeter
### BEGIN HIDDEN TESTS
assert rect2_perimeter == 24
### END HIDDEN TESTS
```

## Submit

Midterm Complete - go ahead and check your work, restart & run all to double check it all runs without issue, and submit on datahub!