# COGS 18 - Practice Exam 1

This is the practice exam for Winter 2023. It covers topics through the Loops Lecture.

This exam is out 0 points and worth 0% of your grade, but the real thing will be out of 12.5 points, worth 12.5% of your grade.

**PLEASE DO NOT CHANGE THE NAME OF THIS FILE.**

**PLEASE DO NOT COPY & PASTE OR DELETE CELLS INLCUDED IN THE ASSIGNMENT.** (Note that you can add additional cells, if you want to test things out.

## Instructions

#### Timing
- The exam is designed to take you ~1h.
- There will be a >24 hour window during which you can complete this exam (due Mon 11:59 PM).

#### The Rules
- You are to complete this exam on your own.
- This is open-notes & open-Google
- You may not talk to any humans about this exam. 
- The following are all *prohibited*:
    - text/phone/online chat communication
    - posting questions to a message board where a human could respond (Campuswire, Discord, Chegg, any similar site)
    - viewing exam questions from a message board (as described above)
    - asking anyone via any form about a question on this test directly


 <span style="color: red;">Note: </span> There _is_ a chance for partial credit on some questions, so _some_ code is better than _no_ code. Even if it throws an error, having some code that partially answers the question will benefit you/your grade.

### Q0 - Honor Code (0.15 points)

In the cell below, include a variable `honor_code` that stores the boolean `True` if you agree to the following statement:

>I agree that this exam was completed individually with the knowlege in my brain, information the notes from this course, and/or with searching for help on the Internet *without searching for answers to the text from these questions directly*. I did not ask anyone about specific questions on this exam. I did not post these questions (in part or in whole) on the Internet. I did not copy answers to these questions from anywhere or anyone else. I understand all code I've written on this exam and could explain my answers to someone else if asked.


<!-- grade_id: cell-9dae831891d7a1ad, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
honor_code = True
### END SOLUTION
```

<!-- grade_id: cell-8a6d1a928512e45c, points: 0.15, grade: true, solution: false -->
```python
assert honor_code
```

## Part 1: Variables & Operators (3.65 points)

### Q1 - Variables (0.75 points)

In the cell below, create three variables: `my_string`, `my_float`, `my_int`, and `my_bool`. 

- `my_string` should store any string of your choice having length 5 and should start with character "s".
- `my_float` should be a decimal point value that is less than 10.
- `my_int` should be a positive integer divisible by 5.
- `my_bool` should be a boolean value.

The specific values you use to generate the variables are up to you, as long as they fulfill the criteria specified.

<!-- grade_id: q1a, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
# actual values stored will differ
my_string = 'sssss'
my_float = 2.4
my_int = 15
my_bool = True

### END SOLUTION
```

<!-- grade_id: q1-t1, points: 0.25, grade: true, solution: false -->
```python
# check variable exists
assert len(my_string) == 5 and (my_string[0] == 's' or my_string[0] == 'S') 
### BEGIN HIDDEN TESTS
assert isinstance(my_float, float) and my_float<10.0
assert isinstance(my_int, int) and my_int%5==0
assert type(my_bool) == bool
### END HIDDEN TESTS
```

### Q2 - Math Operators (1 point)

Use the variables provided in the cell below (`var_low`, `var_mid` and `var_high`) to generate the variables: `var_150` and `var_14` according to the following specifications:

- `ans_1` | compute the answer when var_low is doubled and multiplied with var_mid and then added to var_high
- `ans_2` | use the 3 variables provided once each and *at least two* of the math operators (`+`, `-`, `*`, `/`) to store the value 75

```python
# these variables provided for you
# this cell is read-only (cannot be changed)
var_low = 2
var_mid = 10
var_high = 15
```

<!-- grade_id: q2a, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
# actual operations will differ
ans_1 = 2*var_low*var_mid + var_high
ans_2 = (var_high * var_mid) / var_low

print(ans_1, ans_2)
### END SOLUTION
```

Output:
```
55 75.0
```

<!-- grade_id: q2-t1, points: 0.35, grade: true, solution: false -->
```python
assert ans_1 is not None
### BEGIN HIDDEN TESTS
assert ans_1 == 55
### END HIDDEN TESTS
```

<!-- grade_id: q2-t2, points: 0.35, grade: true, solution: false -->
```python
assert ans_2 is not None
### BEGIN HIDDEN TESTS
assert ans_2 == 75.0
### END HIDDEN TESTS
```

```python

```

**This "blank" cell included intentionally.** Do not do anything here. (It's being used in grading.)

=== BEGIN MARK SCHEME ===

- correct operators used (0.15)
- variables referenced; did not hard-code numbers (0.15)  

Note: It's ok if they also referenced additional values other than the variables provided.

=== END MARK SCHEME ===


### Q3 - Comparison Operators (1 point)

Store your favorite number in the variable `fav_num` (can be any number). 

Then, store another number between -100 and 100 (both not inclusive) in the variable `comparison_num`

Finally, referencing both `fav_num` and `comparison_num` *and* using at least one comparison operator, define two variables: `true_comp` and `false_comp`

These variables must satisfy the following requirements:

1. Each variable must be defined using one comparison operator (`<`,`>`,`<=`,`>=`,`!=`,`==`).
2. Each variable must use and refer to `fav_num` and `comparison_num` (and no other values or variables)
4. `true_comp` should evaluate as the boolean `False`; `false_comp` should evaluate as the boolean `True`

<!-- grade_id: q3a, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
fav_num = 60 # specific values will differ
comparison_num = -10

true_comp = fav_num < comparison_num
false_comp = fav_num > comparison_num
### END SOLUTION
```

<!-- grade_id: q3-t1, points: 0.25, grade: true, solution: false -->
```python
assert true_comp is not None
### BEGIN HIDDEN TESTS
assert not true_comp
### END HIDDEN TESTS
```

<!-- grade_id: q3-t2, points: 0.25, grade: true, solution: false -->
```python
assert false_comp is not None
### BEGIN HIDDEN TESTS
assert false_comp
### END HIDDEN TESTS
```

```python

```

**This "blank" cell included intentionally.** Do not do anything here. (It's being used in grading.)

=== BEGIN MARK SCHEME ===

- both variables reference `fav_num` and `comparison_num` (**0.25 pts**)
- both variables reference use at least one comparison operator (**0.25 pts**)

=== END MARK SCHEME ===

### Q4 - Membership Operators (0.4 points)

Define an *string* variable of length 10 called `my_string` (all capital letters), such that when you check whether 'COGS' is a member of `my_string` it returns `True`.

<!-- grade_id: cell-f4e64b6777aff0a4, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
# students could define *either* a string OR a tuple here
my_string = 'COGSCOGS18'
### END SOLUTION
```

<!-- grade_id: cell-92c7c3227a68644b, points: 0.1, grade: true, solution: false -->
```python
assert my_string is not None
### BEGIN HIDDEN TESTS
assert type(my_string) is str
### END HIDDEN TESTS
```

<!-- grade_id: cell-3be15b8b3f026f5e, points: 0.1, grade: true, solution: false -->
```python
# hidden tests to check length
### BEGIN HIDDEN TESTS
assert len(my_string) == 10
### END HIDDEN TESTS
```

<!-- grade_id: cell-5c81a4b7ed540a90, points: 0.1, grade: true, solution: false -->
```python
# hidden tests to check conditions specified in instructions
### BEGIN HIDDEN TESTS
assert 'COGS' in my_string
### END HIDDEN TESTS
```

<!-- grade_id: cell-f77a2729e8a01ea4, points: 0.1, grade: true, solution: false -->
```python
# hidden tests to check conditions specified in instructions
### BEGIN HIDDEN TESTS
assert my_string.isupper()
### END HIDDEN TESTS
```

## Part 2: Functions + Conditionals  (6.2 points)

### Q5 - Conditionals (1.6 points)

You're in-charge of safety of rides and following protocols in an amusement park. 

For this you have to consider both *whether* the person is an adult or a child, depending on their age criterion and also depending on the danger level of the ride. If the person should receive a ticket, the variable `ticket` should store the boolean `True`; otherwise, `ticket` should store `False`.

More specifically, you have to write a conditional that references the three variables provided below and considers the following logic:

It should specify that a ticket should be provided under the following three conditions:
1. If danger level of ride is 2 then the person must be an adult (atleast 18 years age)
2. If danger level of ride is 1 then the person must be an adult (atleast 18 years age) or the person can be greater than 14 years of age but must be accompanied by an adult
3. If danger level of ride is 0, any person of any age can get on the ride

Otherwise, the person should not be provided a ticket.


```python
# these variables provided for you
# this cell is read-only (cannot be changed)
danger_level = 1
accompanied_by_adult = True
```

<!-- grade_id: q6a, points: 0, grade: false, solution: true -->
```python
age = 15

### BEGIN SOLUTION
if danger_level == 0 and age > 0:
    ticket = True
elif danger_level == 1 and ((age >= 18) or (age>14 and accompanied_by_adult)):
    ticket = True
elif danger_level == 2 and age >= 18:
    ticket = True
else:
    ticket = False

### END SOLUTION
```

<!-- grade_id: q6-t1, points: 0.6, grade: true, solution: false -->
```python
assert ticket is not None
### BEGIN HIDDEN TESTS
assert ticket
### END HIDDEN TESTS
```

```python

```

**This "blank" cell included intentionally.** Do not do anything here. (It's being used in grading.)


=== BEGIN MARK SCHEME ===

- uses `if`/`elif`/`elif`/`else` syntax correctly; half credit if `if`/`if`/`if`/`else` used (**0.4 points**)
- conditionals correct 
    - references variables (`private_mph`, `local_mph`, `hwy_mph`) and doesn't hard-code (**0.2 points**)
    - references `zone` (**0.1 points**)
    - uses `>` and `and`/`&` correctly  (**0.2 points**)
- `ticket` assigned a boolean value (even if logic is incorrect) (**0.1 points**)

=== END MARK SCHEME ===

### Q6 - Function Execution (2 points)


Here is a function we will talk about below.

MAKE SURE YOU EXECUTE THE FUNCTION DEFINTION IN THE CELL BELOW.

```python
# this function is provided for you
# this cell is read-only (cannot be changed)
def introduce(name, hobby, food):
    
    output = 'Hi! my name is ' + name + '. My hobby is ' + hobby + ' and my favourite food is ' + food + '.'
    
    return output
```

Use the cell below to print the output of 2 calls to the function above using 2 different people.

Your solution should use the function defined above and print the output of the function doing it for 2 different people.

```python
### BEGIN SOLUTION
print(introduce('Mike','guitar','taco'))
print(introduce('Emily','singing','sushi'))
### END SOLUTION
```

Output:
```
Hi! my name is Mike. My hobby is guitar and my favourite food is taco.
Hi! my name is Emily. My hobby is singing and my favourite food is sushi.
```

**This "blank" cell included intentionally.** Do not do anything here. (It's being used in grading.)

=== BEGIN MARK SCHEME ===
- `output_list` initialized with a value [] (**0.25 points**)
- `for` loop thorugh `cars`/`zones`/range of len (**0.5 points**)
- indexed to get value from list correctly (**0.5 points**)
- conditionals included (not grading for correctness, as last question did) (**0.25 points**)
- list appended correctly (**0.25 points**)
    
=== END MARK SCHEME ===

### Q7 - Functions + Conditionals (2 points)

Background: Students are often curious about their grade in classes. Here, you're going to help students figure out their grade by writing a function.

You need to define a function called "check_grade" by taking in one parameter called score, attendance.

The status of a student will only be shown if attendance is greater than 75, otherwise your grade would be displayed as 'F'.
Considering that a student has their attendance greater than 75, 'A' grade is given to students who have a score of 90+, 'B' grade is  given to students who have a score of 75-90, 'C' grade is  given to students who have a score of 60-75, and if less than 60, 'D' grade is given.

Find the output, of a student (grade of that student) given the attendance and the score.


<!-- grade_id: cell-1ea6cb63f391a4cf, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def check_grade(score, attendance):
    if attendance < 75:
        grade = 'F'
    else:
        if score > 90:
            grade = 'A'
        elif score >75:
            grade = 'B'
        elif score >60:
            grade = 'C'
        else:
            grade = 'D'
    
    return grade
### END SOLUTION
```

<!-- grade_id: cell-b55e3349d00c555e, points: 0.6, grade: true, solution: false -->
```python

### BEGIN HIDDEN TESTS
grade = check_grade(100, 60)
assert grade == 'F'

grade = check_grade(20, 75)
assert grade == 'D'


grade = check_grade(91, 93)
assert grade == 'A'


grade = check_grade(87, 90)
assert grade == 'B'


grade = check_grade(70, 89)
assert grade == 'C'
### END HIDDEN TESTS
```

**This "blank" cell included intentionally.** Do not do anything here. (It's being used in grading.)

=== BEGIN MARK SCHEME ===
- `prof_hours` initialized (**0.25 points**)
- `for` loop set up correctly (**0.25 points**)
- conditional correct (references long_tasks; short_tasks) (**0.25 points**) 
- `prof_hours` incremented correctly (**0.5 points**)  

=== END MARK SCHEME ===

## Part 3: Collections and Indexing (2.5 points)

### Q8 - Lists & Indexing (1 point)

The list `burritos` has been provided in the cell below for you.

***Use the `burritos` variable and indexing*** to generate each of the four variables at left below, so that each stores the output at right below. For example, `burritos_a` should refer to `burritos` in its answer and use indexing to store the string 'Cali'.

- `burritos_a` | `'Chile Verde'`
- `burritos_b` | `['Bean and Cheese', 'Steak Ranchero', 'Carnitas']`
- `burritos_c` | `['Adobada', 'Carnitas']`
- `burritos_d` | `['Chicken', 'Cali', 'Steak Ranchero']`

```python
# this variables provided for you
# this cell is read-only (cannot be changed)
burritos = ['Chicken', 'Carne Asada', 'Chile Verde', 'Cali', 
            'Adobada', 'Bean and Cheese', 'Steak Ranchero',
            'Carnitas']
```

<!-- grade_id: cell-2f6afe8e2f00389a, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
# actual values will differ
burritos_a = burritos[2]
burritos_b = burritos[-3:]
burritos_c = burritos[4::3]
burritos_d = burritos[::3]
### END SOLUTION
```

<!-- grade_id: cell-b087dcf49a19a876, points: 0.25, grade: true, solution: false -->
```python
assert burritos_a
### BEGIN HIDDEN TESTS
assert burritos_a == 'Chile Verde'
### END HIDDEN TESTS
```

<!-- grade_id: cell-db3ae047ed3001b3, points: 0.25, grade: true, solution: false -->
```python
assert burritos_b
### BEGIN HIDDEN TESTS
assert burritos_b == ['Bean and Cheese', 'Steak Ranchero', 'Carnitas']
### END HIDDEN TESTS
```

<!-- grade_id: cell-9c6b23230d2dd7e2, points: 0.25, grade: true, solution: false -->
```python
assert burritos_c
### BEGIN HIDDEN TESTS
assert burritos_c == ['Adobada', 'Carnitas']
### END HIDDEN TESTS
```

<!-- grade_id: cell-a0ae2d0f54500a1d, points: 0.25, grade: true, solution: false -->
```python
assert burritos_d
### BEGIN HIDDEN TESTS
assert burritos_d == ['Chicken', 'Cali', 'Steak Ranchero']
### END HIDDEN TESTS
```

**This "blank" cell included intentionally.** Do not do anything here. (It's being used in grading.)

=== BEGIN MARK SCHEME ===

- if hard-coded (doesn't use indexing & reference `burritos`), remove autograder credit

=== END MARK SCHEME ===


### Q9 - Tuples (1 point)

In the cell below, describe a situation where you would store information in a tuple (give an example). Be clear about the information or values you would store in a tuple and why that variable type would be appropriate for your situation. 

Note: you do not *have* to write any python code here; we're looking for a written explanation.

<!-- grade_id: cell-de38aa298e12df9e, points: 0, grade: true, solution: true -->


**This "blank" cell included intentionally.** Do not do anything here. (It's being used in grading.)

=== BEGIN MARK SCHEME ===

- example provided includes multiple pieces of information being stored (**0.25 pts**)
- information being stored would reasonably NOT want to be changed (**0.50 pts**)
- student explanation links immutability to their example (**0.25 pts**)

=== END MARK SCHEME ===

### Q10 - Collections: Indexing & Mutating (0.5 point)

One of the TAs in COGS 18 is trying to learn more about the colleges at UCSD, so she's pulled together a list of the colleges. Unfortunately, she's got some bad information. Using code constructs discussed in class, replace any element in `ucsd_colleges` that is NOT a college at UCSD with the string `ERROR`. Your updated list should be stored in the same variable: `ucsd_colleges`.

Note: Prof Ellis is NOT trying to trip you up on spelling or anything here, so you don't have to check each college's spelling or anything like that.

```python
# this variable provided for you
# this cell is read-only (cannot be changed)
ucsd_colleges = ['Revelle',
                 'Sarah Marshall',
                 'John Muir',
                 'Thurgood Marshall',
                 'Earl Warren',
                 'Eleanor Roosevelt',
                 'Torrey Pines',
                 'Sixth',
                 'Seventh']
```

<!-- grade_id: cell-27ed0de1ff49cfac, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
ucsd_colleges[1] = 'ERROR'
ucsd_colleges[6] = 'ERROR'
### END SOLUTION
```

<!-- grade_id: cell-ecac9c070cf94edf, points: 0.25, grade: true, solution: false -->
```python
assert ucsd_colleges is not None
### BEGIN HIDDEN TESTS
assert ucsd_colleges[1] == 'ERROR'
### END HIDDEN TESTS
```

<!-- grade_id: cell-e6240eecb22337ae, points: 0.25, grade: true, solution: false -->
```python
# hidden test
### BEGIN HIDDEN TESTS
assert ucsd_colleges[6] == 'ERROR'
### END HIDDEN TESTS
```

**This "blank" cell included intentionally.** Do not do anything here. (It's being used in grading.)

=== BEGIN MARK SCHEME ===

- if hard-coded (doesn't use indexing/conditionals & does not reference `ucsd_colleges`), remove autograder credit

=== END MARK SCHEME ===


### Submit on Datahub

Good work - you're done!

Check your work, and then when you're ready, **submit on datahub**.

We will not have your exam until you click submit and see this show up under 'submitted assignments'.

```python

```