# COGS 18 - Exam 1 (Take-Home)

This is the take-home portion of the first exam for Spring 2024. It covers the technical implementation of topics through the Collections Lecture.

This take-home portion of the exam is out of 2.5 points, worth 2.5% of your grade.

**PLEASE DO NOT CHANGE THE NAME OF THIS FILE.**

**PLEASE DO NOT COPY & PASTE OR DELETE CELLS INLCUDED IN THE ASSIGNMENT.** (Note that you can add additional cells, if you want to test things out.

**In-class** answer key [here](https://docs.google.com/document/d/1iN60zQJ2FS0fgWVKVTZZIqihc6JSfrOP6DHd6xDDXDE/edit?usp=sharing)

#### Instructions

#### Timing
- There take-home portion will be released 3:30 the day of the exam and due the same day at 11:59PM

#### The Rules
- You are to complete this exam on your own.
- This is open-notes & open-Internet.
- You may not talk to any humans about this exam. 
- Clarification questions will ***not*** be allowed and there will be no posting on Piazza about the exam at all. 
     - If you are confused about wording, add a note to your exam explaining your confusion and how you interpreted the question. 
    - If you have a technical issue completing the exam, message (Piazza or email) Prof Ellis ASAP.

 <span style="color: red;">Note: </span> This portion of the midterm will be autograded, similar to an assignment. `assert` statements are there to guide your thinking but do not guarantee you've gotten the answer correct.

### Q0 - Honor Code (0.05 points)

In the cell below, include a variable `honor_code` that stores the boolean `True` if you agree to the following statement:

>I agree that this exam was completed individually with the knowlege in my brain, information the notes from this course, and/or with searching for help on the Internet *without searching for answers to the text from these questions directly*. I did not ask anyone about specific questions on this exam. I did not post these questions (in part or in whole) on the Internet. I did not copy answers to these questions from anywhere or anyone else. I understand all code I've written on this exam and could explain my answers to someone else if asked.


<!-- grade_id: cell-9dae831891d7a1ad, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
honor_code = True
### END SOLUTION
```

<!-- grade_id: cell-8a6d1a928512e45c, points: 0.05, grade: true, solution: false -->
```python
assert honor_code
```

### Background


For these next two questions, each question will be related to the same task: writing code to determine if a given year is a leap year. The idea here is that there will be two functions: 1) a function that determines whether or not a given year is a leap year 2) a function that displays the output in a sentence that conveys the year and whether or not it is a leap year.




### Q1 - Debugging a function (1 pt)

On the in-person portion of the midterm you were told the following: You sit down to write a function (`is_leap_year`) that will return `True` if the input year is a leap year and `False` if not.  However, when you sit down to write your function, you draft the following:


```python
def is_leap_year(year):
    if year / 4 = 0:
        if year / 100 = 0:
            if year / 400 = 0:
                output =  True
            else:
                output = False
        else:
            output = True
    else:
        output = False
```

You were also provided the following information: 


The following rules apply for determining whether something is a leap year:

- Every year that is evenly divisible by 4 is a leap year, except for:
    - Years that are evenly divisible by 100 are not leap years, unless:
        - The year is also evenly divisible by 400, in which case it is a leap year.
        
With this in mind, you know that `is_leap_year(2024)` should return `True` where `is_leap_year(2023)` should return `False`. But it's not doing that...yet. **Debug this function so that it accomplishes what you had intended.** Do not change the name of the function or the input parameter name.

<!-- grade_id: cell-890ee14a107fa8f5, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def is_leap_year(year):
    
    if year % 4 == 0:
        if year % 100 == 0:
            if year % 400 == 0:
                output =  True
            else:
                output = False
        else:
            output = True
    else:
        output = False
    
    return output
### END SOLUTION
```

```python
# you can use this cell to test/execute/check your thinking (optional)
```

<!-- grade_id: cell-986ea16845831923, points: 0.25, grade: true, solution: false -->
```python
# check no typo in function name
assert callable(is_leap_year)
```

<!-- grade_id: cell-0ad5a91264167d53, points: 0.25, grade: true, solution: false -->
```python
# hidden test that checks the output is of the expected type
### BEGIN HIDDEN TESTS
# note: this would still give credit, even if you had an error in Q1
assert type(is_leap_year(2022)) == bool 
### END HIDDEN TESTS
```

<!-- grade_id: cell-7d903b9bfc07d385, points: 0.25, grade: true, solution: false -->
```python
# hidden tests that check the function works with provided examples
### BEGIN HIDDEN TESTS
assert is_leap_year(2024)
assert not is_leap_year(2023)
### END HIDDEN TESTS
```

<!-- grade_id: cell-0693f283f6eaafeb, points: 0.25, grade: true, solution: false -->
```python
# hidden tests that check the function works as specified beyond provided examples
### BEGIN HIDDEN TESTS
assert is_leap_year(300) == False
assert is_leap_year(400) == True
assert is_leap_year(32) == True
assert is_leap_year(401) == False
### END HIDDEN TESTS

```

### Q2 - Defining a function (1.45 pts)

After you get your `is_leap_year()` function working as intended, you want to write a function `output_leap_year()`. This function should take in a `year` as input, determine if the `year` is a leap year, and, if it is, output '{year} is a leap year!' or, if it is not, '{year} is not a leap year!'...where {year} is replaced with the year input. 

Note: Spacing and punctuation in your string matter. For example: 
- `output_leap_year(2023)` should return the string: `'2023 is not a leap year!'` 
- `output_leap_year(2024)` would return `'2024 is a leap year!'`

<!-- grade_id: cell-5406a3284790c7de, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def output_leap_year(year):
    
    yes_leap_year = is_leap_year(year)
    
    if yes_leap_year:
        output = str(year) + ' is a leap year!'
    elif not yes_leap_year:
        output = str(year) + ' is not a leap year!'
    else:
        output = None
    
    return output
### END SOLUTION
```

```python
# you can use this cell to test/execute/check your thinking (optional)
```

<!-- grade_id: cell-b6f17c05f4455dff, points: 0.25, grade: true, solution: false -->
```python
assert callable(output_leap_year)
### BEGIN HIDDEN TESTS
# this ensures you have the correct is_leap_year function
# so that you're able to get full credit below...even if you didn't figure out Q1
def is_leap_year(year):
    
    if year % 4 == 0:
        if year % 100 == 0:
            if year % 400 == 0:
                output =  True
            else:
                output = False
        else:
            output = True
    else:
        output = False
    
    return output
### END HIDDEN TESTS
```

<!-- grade_id: cell-3ed4c4555bf2aad6, points: 0.25, grade: true, solution: false -->
```python
# hidden test that checks the output is of the expected type
### BEGIN HIDDEN TESTS
# note: this would still give credit, even if you had an error in Q1
assert type(output_leap_year(2022)) == str 
### END HIDDEN TESTS
```

<!-- grade_id: cell-1ce6db55cfb0fe1a, points: 0.25, grade: true, solution: false -->
```python
# hidden tests that check the parameter name
### BEGIN HIDDEN TESTS
# note: this would still give credit, even if you had an error in Q1
assert output_leap_year(year=2022)
### END HIDDEN TESTS
```

<!-- grade_id: cell-1ff39fd3a0967d23, points: 0.25, grade: true, solution: false -->
```python
# hidden tests that check the provided examples
### BEGIN HIDDEN TESTS
# note: this would still give credit, even if you had an error in Q1
assert output_leap_year(2023) == '2023 is not a leap year!'
assert output_leap_year(2024) == '2024 is a leap year!'
### END HIDDEN TESTS
```

<!-- grade_id: cell-0cdb562a88b1d7d6, points: 0.45, grade: true, solution: false -->
```python
# hidden tests that check overall functionality beyond provided cases
### BEGIN HIDDEN TESTS
# note: this would still give credit, even if you had an error in Q1
assert output_leap_year(300) == '300 is not a leap year!'
assert output_leap_year(400) == '400 is a leap year!'
assert output_leap_year(32) == '32 is a leap year!'
assert output_leap_year(401) == '401 is not a leap year!'
### END HIDDEN TESTS
```